---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1

```yaml
type: NormalExercise
key: b2353fb483
lang: r
xp: 100
skills: 1
```

This is the assignment text. It should help provide students with the background information needed.
The instructions that follow should be in bullet point form with clear guidance for what is expected.

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. It should get students 50% of the way to the correct answer.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Video 1

```yaml
type: VideoExercise
key: a2325729b8
xp: 50
```

`@projector_key`
00b4da47cafa4bbe9f11e8b8de7a8fd4

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 0f5f3a6d6e
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
